<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 2 quick review: research questions</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Consider this RQ:</p><blockquote><p>In elite female netball players, do players in defence positions have the same average number of knee injuries (per player, per season) compared to players in attacking positions?</p></blockquote><p>Are the following statements <em>true</em> or <em>false</em>?</p><ol><li><p>The <em>comparison</em> is &lsquo;between knee injuries and other types of injuries&rsquo;. [[input:ans1]] [[validation:ans1]]</p></li><li><p>The <em>comparison</em> in this RQ is a <em>between</em>-individuals comparison. [[input:ans2]] [[validation:ans2]]</p></li><li><p>The <em>outcome</em> is &lsquo;the average number of knee injuries per player, per season&rsquo;. [[input:ans3]] [[validation:ans3]]</p></li><li><p>The <em>response variable</em> is &lsquo;the average number of knee injuries per season&rsquo;. [[input:ans4]] [[validation:ans4]]</p></li><li><p>The <em>unit of analysis</em> is &lsquo;the number of knee injuries&rsquo;. [[input:ans5]] [[validation:ans5]]</p></li><li><p>The <em>unit of observation</em> is &lsquo;the elite netball player&rsquo;. [[input:ans6]] [[validation:ans6]]</p></li><li><p>This RQ is a descriptive RQ. [[input:ans7]] [[validation:ans7]]</p></li><li><p>This RQ is an estimation-type RQ. [[input:ans8]] [[validation:ans8]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>For this RQ: the <em>population</em> is elite female netball players; the <em>outcome</em> is the average number of knee injuries per player per season; the <em>comparison</em> is between defence and attacking positions, a between-individuals comparison; there is no <em>intervention</em>. The response variable is the number of knee injuries per player per season, recorded on each player (the unit of observation and the unit of analysis). Because the RQ asks whether the two averages are the same, it is a two-tailed, decision-making, relational RQ.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 2 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans6</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans7</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans8</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The comparison is between the two playing positions (defence and attack); knee injuries are what is measured, not what is compared.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The comparison is between the two playing positions (defence and attack); knee injuries are what is measured, not what is compared.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Defenders and attackers are different players, so the two groups being compared are made of different individuals.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Defenders and attackers are different players, so the two groups being compared are made of different individuals.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The outcome is the numerical summary (here, an average) of the response variable over many individuals.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The outcome is the numerical summary (here, an average) of the response variable over many individuals.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The response variable is recorded on each individual: the number of knee injuries a player has in a season. The average of those values is the outcome.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The response variable is recorded on each individual: the number of knee injuries a player has in a season. The average of those values is the outcome.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The number of knee injuries is a variable. The unit of analysis is the smallest independent example of the population, here the individual player.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The number of knee injuries is a variable. The unit of analysis is the smallest independent example of the population, here the individual player.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans6</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans6</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans6-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The observations (knee injuries) are made on each player, so the player is the unit of observation.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans6-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The observations (knee injuries) are made on each player, so the player is the unit of observation.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans7</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans7</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans7-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A descriptive RQ studies one outcome in one population. This RQ compares two groups, so it is a relational RQ.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans7-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A descriptive RQ studies one outcome in one population. This RQ compares two groups, so it is a relational RQ.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans8</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans8</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans8-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The RQ asks whether the averages are the same or not, a yes-or-no question, so it is a decision-making RQ, not an estimation RQ.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans8-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The RQ asks whether the averages are the same or not, a yes-or-no question, so it is a decision-making RQ, not an estimation RQ.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>11</testcase>
<description>ans6: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>12</testcase>
<description>ans6: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>13</testcase>
<description>ans7: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>14</testcase>
<description>ans7: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>15</testcase>
<description>ans8: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>16</testcase>
<description>ans8: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>