<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 3 quick review: overview of research design</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Martnes and Bere (2023) compared the average time to complete a journey when (p. 1)</p><blockquote><p>&hellip; riding an electric-assisted bicycle with cargo (30&nbsp;kg) and without cargo&hellip;</p></blockquote><p>They recorded the age, height, weight, and resting metabolic rate of all subjects who completed the 4.5&nbsp;km ride. Each subject was allocated to ride both with <em>and</em> without cargo.</p><p>Are the following statements <em>true</em> or <em>false</em>?</p><ol><li><p>The explanatory variable is &lsquo;the age of the subjects&rsquo;. [[input:ans1]] [[validation:ans1]]</p></li><li><p>&lsquo;The height of the subjects&rsquo; is a lurking variable. [[input:ans2]] [[validation:ans2]]</p></li><li><p>The explanatory variable is &lsquo;whether the bicycle is ridden with or without cargo&rsquo;. [[input:ans3]] [[validation:ans3]]</p></li><li><p>&lsquo;Weight&rsquo; is an extraneous variable. [[input:ans4]] [[validation:ans4]]</p></li><li><p>The response variable is &lsquo;the time to complete the journey&rsquo;. [[input:ans5]] [[validation:ans5]]</p></li><li><p>&lsquo;Age&rsquo; is a possible confounding variable. [[input:ans6]] [[validation:ans6]]</p></li><li><p>&lsquo;Resting metabolic rate&rsquo; is a possible confounding variable. [[input:ans7]] [[validation:ans7]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>The <em>response variable</em> is the time to complete the journey; the <em>explanatory variable</em> is whether the bicycle is ridden with or without cargo. Age, height, weight and resting metabolic rate are <em>extraneous</em> variables; because they were recorded they are not lurking variables. Because each subject rode both with and without cargo (a within-individuals comparison), none of these can be <em>confounding</em>: a confounding variable has to be related to the explanatory variable as well as to the response, and here every subject appears in both conditions.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 3 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans6</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans7</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The explanatory variable is the one whose values are being compared: riding with or without cargo. Age was recorded, but it is an extraneous variable.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The explanatory variable is the one whose values are being compared: riding with or without cargo. Age was recorded, but it is an extraneous variable.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A lurking variable is an extraneous variable that was <em>not</em> recorded. Height was recorded, so it is an extraneous variable, not a lurking one.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A lurking variable is an extraneous variable that was <em>not</em> recorded. Height was recorded, so it is an extraneous variable, not a lurking one.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The study compares journey times with and without cargo, so with-or-without cargo is the explanatory variable.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The study compares journey times with and without cargo, so with-or-without cargo is the explanatory variable.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Weight is not the variable of interest but could influence the journey time, so it is an extraneous variable (a recorded one).</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Weight is not the variable of interest but could influence the journey time, so it is an extraneous variable (a recorded one).</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The journey time is what was measured on each ride and compared between the two conditions, so it is the response variable.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The journey time is what was measured on each ride and compared between the two conditions, so it is the response variable.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans6</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans6</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans6-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A confounding variable must be related to both the response and the explanatory variable. Every subject rode both with and without cargo, so age is the same in both conditions and cannot be confounding.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans6-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A confounding variable must be related to both the response and the explanatory variable. Every subject rode both with and without cargo, so age is the same in both conditions and cannot be confounding.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans7</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans7</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans7-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. Each subject rode under both conditions, so resting metabolic rate cannot differ between the with-cargo and without-cargo rides; it is extraneous, not confounding.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans7-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. Each subject rode under both conditions, so resting metabolic rate cannot differ between the with-cargo and without-cargo rides; it is extraneous, not confounding.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>11</testcase>
<description>ans6: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>12</testcase>
<description>ans6: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>13</testcase>
<description>ans7: correct (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>14</testcase>
<description>ans7: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>