<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 5 quick review: ethics in research</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Are the following statements <em>true</em> or <em>false</em>?</p><ol><li><p>Ethics apply for <em>any</em> type of study. [[input:ans1]] [[validation:ans1]]</p></li><li><p>Ethics only refer to the interactions of the researchers with participants in the study. [[input:ans2]] [[validation:ans2]]</p></li><li><p>Ethics only apply when <em>people</em> are the individuals. [[input:ans3]] [[validation:ans3]]</p></li><li><p>Ethics only apply when <em>people</em> or <em>animals</em> are the individuals. [[input:ans4]] [[validation:ans4]]</p></li><li><p>Ethics can extend to storage of data. [[input:ans5]] [[validation:ans5]]</p></li><li><p>Ethics only apply to the design of the study. [[input:ans6]] [[validation:ans6]]</p></li><li><p>Ethics apply even to the analysis of the data. [[input:ans7]] [[validation:ans7]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>Studies must be ethical, and any formal study must obtain ethical approval before beginning. Ethics covers far more than the treatment of participants: acknowledgements, analysis methods, confidentiality, consent, economic, environmental, legal, physical, psychological and social risks, funding, incentives, plagiarism, resourcing, sample size, and the storage of data. It applies to every type of study and to every stage of the research process.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 5 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans6</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans7</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Every study, whether descriptive, observational or experimental, must be ethical and any formal study needs ethical approval before it begins.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Every study, whether descriptive, observational or experimental, must be ethical and any formal study needs ethical approval before it begins.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. Ethics also covers the analysis, the storage of data, plagiarism, acknowledgements, funding, and environmental, legal and economic risks, not only contact with participants.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. Ethics also covers the analysis, the storage of data, plagiarism, acknowledgements, funding, and environmental, legal and economic risks, not only contact with participants.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. Studies of animals raise ethical issues too, and even studies of objects or the environment carry ethical duties about data, honesty and risks.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. Studies of animals raise ethical issues too, and even studies of objects or the environment carry ethical duties about data, honesty and risks.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. Ethical issues such as environmental risks, honest analysis and reporting, plagiarism and data storage apply to any study, whatever the individuals are.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. Ethical issues such as environmental risks, honest analysis and reporting, plagiarism and data storage apply to any study, whatever the individuals are.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. How data are stored, for how long, and who can access them (confidentiality) are ethical matters listed in the chapter.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. How data are stored, for how long, and who can access them (confidentiality) are ethical matters listed in the chapter.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans6</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans6</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans6-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. Ethics applies through the whole research process: design, data collection, analysis, storage and reporting.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans6-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. Ethics applies through the whole research process: design, data collection, analysis, storage and reporting.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans7</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans7</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans7-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Choosing analysis methods honestly, and not manipulating or selectively reporting results, are ethical obligations.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans7-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Choosing analysis methods honestly, and not manipulating or selectively reporting results, are ethical obligations.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>11</testcase>
<description>ans6: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>12</testcase>
<description>ans6: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>13</testcase>
<description>ans7: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>14</testcase>
<description>ans7: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>