<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 7 quick review: internal validity</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Doosti-Irani et al. (2016) wanted to determine the relationship between the depth of bruising on apples and the impact force. The researchers purposefully hit apples with three different <em>forces</em> (200, 700 and 1200&nbsp;mJ) to inflict bruises on the apples. The researchers then recorded the <em>depth</em> of the bruising. The study was conducted separately for three different <em>regions</em> of the apple (lower; middle; upper), and each apple was only used once.</p><p>Are the following statements <em>true</em> or <em>false</em>?</p><ol><li><p>The <em>response variable</em> is &lsquo;the depth of bruising&rsquo;. [[input:ans1]] [[validation:ans1]]</p></li><li><p>The <em>explanatory variable</em> is &lsquo;the force used on the apples&rsquo;. [[input:ans2]] [[validation:ans2]]</p></li><li><p>The variable &lsquo;location of the bruising&rsquo; would be classified as a confounding variable. [[input:ans3]] [[validation:ans3]]</p></li><li><p>The researchers could minimise the effects of confounding by incorporating potential confounding variables in the analysis. [[input:ans4]] [[validation:ans4]]</p></li><li><p>The researchers could use random allocation of the treatments to the apples to minimise confounding. [[input:ans5]] [[validation:ans5]]</p></li><li><p>The <em>carryover</em> effect is likely to be a big problem in this study. [[input:ans6]] [[validation:ans6]]</p></li><li><p>The <em>Hawthorne</em> effect is likely to be a big problem in this study. [[input:ans7]] [[validation:ans7]]</p></li><li><p>The <em>placebo</em> effect is likely to be a big problem in this study. [[input:ans8]] [[validation:ans8]]</p></li><li><p>The <em>observer</em> effect is likely to be a big problem in this study. [[input:ans9]] [[validation:ans9]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>The <em>response variable</em> is the depth of bruising and the <em>explanatory variable</em> is the impact force. The region of the apple is an <em>extraneous</em> variable, managed by studying each region separately. Confounding can be managed by restricting, blocking, adjusting in the analysis, and (in experimental studies like this one) random allocation. The Hawthorne, placebo and observer effects concern the behaviour and expectations of people; with apples and an objective measurement they are not a concern, and because every apple was used once there is no carryover effect.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 7 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans6</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans7</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans8</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans9</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The depth of bruising is what was measured on each apple as the result of the impact, so it is the response variable.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The depth of bruising is what was measured on each apple as the result of the impact, so it is the response variable.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The researchers chose and applied the three forces, and compared the bruising between them, so force is the explanatory variable.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The researchers chose and applied the three forces, and compared the bruising between them, so force is the explanatory variable.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The region of the apple is an <em>extraneous</em> variable. The researchers managed it by studying each region separately (blocking), so it cannot be confounded with the force.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The region of the apple is an <em>extraneous</em> variable. The researchers managed it by studying each region separately (blocking), so it cannot be confounded with the force.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Including potential confounders in the analysis is one of the four ways the chapter gives for managing confounding, alongside restricting, blocking and random allocation.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Including potential confounders in the analysis is one of the four ways the chapter gives for managing confounding, alongside restricting, blocking and random allocation.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. This is an experimental study, so the researchers can allocate the three forces to apples at random, which spreads any extraneous differences evenly across the forces.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. This is an experimental study, so the researchers can allocate the three forces to apples at random, which spreads any extraneous differences evenly across the forces.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans6</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans6</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans6-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A carryover effect needs the same individual to receive more than one treatment. Each apple was used only once, so nothing can carry over.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans6-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A carryover effect needs the same individual to receive more than one treatment. Each apple was used only once, so nothing can carry over.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans7</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans7</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans7-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The Hawthorne effect is a change in behaviour because individuals know they are being studied. Apples do not change their behaviour.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans7-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The Hawthorne effect is a change in behaviour because individuals know they are being studied. Apples do not change their behaviour.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans8</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans8</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans8-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A placebo effect comes from an individual's expectation of a treatment. Apples have no expectations.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans8-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A placebo effect comes from an individual's expectation of a treatment. Apples have no expectations.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans9</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans9</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans9-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The depth of a bruise is an objective measurement, so the researchers' expectations are unlikely to influence what is recorded.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans9-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The depth of a bruise is an objective measurement, so the researchers' expectations are unlikely to influence what is recorded.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>11</testcase>
<description>ans6: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>12</testcase>
<description>ans6: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>13</testcase>
<description>ans7: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>14</testcase>
<description>ans7: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>15</testcase>
<description>ans8: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>16</testcase>
<description>ans8: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>17</testcase>
<description>ans9: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans9</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans9-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>18</testcase>
<description>ans9: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans9</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans9-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>