<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 9 quick review: collecting data</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Answer the following questions about questionnaire questions and protocols.</p><ol><li><p>What is the biggest problem with this question: &lsquo;Do you have bromodosis?&rsquo;</p><p>[[input:ans1]] [[validation:ans1]]</p></li><li><p>What is the biggest problem with this question: &lsquo;Do you spend too much time connected to the internet?&rsquo;</p><p>[[input:ans2]] [[validation:ans2]]</p></li><li><p>What is the biggest problem with this question: &lsquo;Do you eat fruits and vegetables?&rsquo;</p><p>[[input:ans3]] [[validation:ans3]]</p></li><li><p><em>True</em> or <em>false</em>: a well-defined protocol allows the researchers to make the study externally valid. [[input:ans4]] [[validation:ans4]]</p></li><li><p><em>True</em> or <em>false</em>: this question is likely to be a <em>leading</em> question. &lsquo;Do you support a ban on drinks sold in unrecyclable plastic bottles?&rsquo; [[input:ans5]] [[validation:ans5]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>Good questionnaire questions avoid four common faults: <em>leading</em> wording that pushes towards an answer; <em>double-barrelled</em> questions that ask two things at once; <em>ambiguous</em> terms such as &lsquo;too much&rsquo; that respondents interpret differently; and <em>language that may not be understood</em>, such as technical or medical words. A written protocol describes exactly how the data are collected so that the study is consistent and reproducible; it supports internal validity, while external validity comes from the sampling.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 9 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>radio</type>
<tans>[[1,false,"It is double-barrelled"],[2,false,"It is a leading question"],[3,true,"It uses language that may not be understood"],[4,false,"It is ambiguous"]]</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>radio</type>
<tans>[[1,false,"It is double-barrelled"],[2,false,"It is a leading question"],[3,false,"It uses language that may not be understood"],[4,true,"It is ambiguous"]]</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>radio</type>
<tans>[[1,true,"It is double-barrelled"],[2,false,"It is a leading question"],[3,false,"It uses language that may not be understood"],[4,false,"It is ambiguous"]]</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>3</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; &lsquo;It uses language that may not be understood&rsquo;. &lsquo;Bromodosis&rsquo; is a medical word for smelly feet; most respondents will not know it, so their answers cannot be trusted.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; the answer is &lsquo;It uses language that may not be understood&rsquo;. &lsquo;Bromodosis&rsquo; is a medical word for smelly feet; most respondents will not know it, so their answers cannot be trusted.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>4</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; &lsquo;It is ambiguous&rsquo;. &lsquo;Too much&rsquo; means something different to every respondent, so the question is ambiguous: two people with the same habits can honestly answer differently.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; the answer is &lsquo;It is ambiguous&rsquo;. &lsquo;Too much&rsquo; means something different to every respondent, so the question is ambiguous: two people with the same habits can honestly answer differently.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>1</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; &lsquo;It is double-barrelled&rsquo;. The question asks two things at once (fruit, and vegetables); someone who eats one but not the other cannot answer it.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; the answer is &lsquo;It is double-barrelled&rsquo;. The question asks two things at once (fruit, and vegetables); someone who eats one but not the other cannot answer it.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. A protocol makes the data collection consistent and reproducible, which protects internal validity. External validity depends on how the sample represents the population, which a protocol cannot fix.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. A protocol makes the data collection consistent and reproducible, which protects internal validity. External validity depends on how the sample represents the population, which a protocol cannot fix.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The word &lsquo;unrecyclable&rsquo; nudges the respondent towards supporting the ban, so the wording leads the answer.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The word &lsquo;unrecyclable&rsquo; nudges the respondent towards supporting the ban, so the wording leads the answer.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (3)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (1)</description>
<testinput>
<name>ans1</name>
<value>1</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (4)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (1)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>1</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (1)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (2)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>2</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (false)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (true)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>3</value>
</testinput>
<testinput>
<name>ans2</name>
<value>4</value>
</testinput>
<testinput>
<name>ans3</name>
<value>1</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>