<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
<name>
<text>Chapter 11 quick review: summarising quantitative data</text>
</name>
<questiontext format="html">
<text><![CDATA[<p>Are the following statements <em>true</em> or <em>false</em>?</p><ol><li><p>The IQR measures the amount of variability in data. [[input:ans1]] [[validation:ans1]]</p></li><li><p>The mean and the median can both be called an &lsquo;average&rsquo;. [[input:ans2]] [[validation:ans2]]</p></li><li><p>The mean and the median are not always the same value. [[input:ans3]] [[validation:ans3]]</p></li><li><p>The range is a simple measure of variation in a set of data. [[input:ans4]] [[validation:ans4]]</p></li><li><p>The standard deviation measures the amount of variability in data. [[input:ans5]] [[validation:ans5]]</p></li><li><p>Another name for the median is \(Q_2\). [[input:ans6]] [[validation:ans6]]</p></li><li><p>\(Q_3\) is the median of the largest half of the data. [[input:ans7]] [[validation:ans7]]</p></li><li><p>The IQR is a useful measure of variation in data that are skewed. [[input:ans8]] [[validation:ans8]]</p></li><li><p>The IQR is the difference between the first and second quartiles. [[input:ans9]] [[validation:ans9]]</p></li><li><p>Another name for the 75th percentile is \(Q_3\). [[input:ans10]] [[validation:ans10]]</p></li><li><p>The units of the standard deviation and the IQR are the same as for the original data. [[input:ans11]] [[validation:ans11]]</p></li></ol>]]></text>
</questiontext>
<questionvariables>
<text></text>
</questionvariables>
<generalfeedback format="html">
<text><![CDATA[<p>Quantitative data are summarised by a measure of the <em>centre</em> (the mean, or the median \(Q_2\)) and a measure of <em>variation</em> (the standard deviation, or the interquartile range \(\text{IQR} = Q_3 - Q_1\), or the simple range). The mean and standard deviation are used when the data are approximately symmetric; the median and IQR when the data are skewed or have outliers, because the quartiles are not affected by extreme values. The standard deviation and the IQR are in the same units as the data.</p>]]></text>
</generalfeedback>
<questionnote format="moodle_auto_format">
<text><![CDATA[Fixed question (chapter 11 quick review), no random variants.]]></text>
</questionnote>
<input>
<name>ans1</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans2</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans3</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans4</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans5</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans6</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans7</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans8</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans9</name>
<type>boolean</type>
<tans>false</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans10</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<input>
<name>ans11</name>
<type>boolean</type>
<tans>true</tans>
<showvalidation>0</showvalidation>
<boxsize>15</boxsize>
<strictsyntax>1</strictsyntax>
<insertstars>0</insertstars>
<syntaxhint></syntaxhint>
<syntaxattribute>0</syntaxattribute>
<forbidwords></forbidwords>
<allowwords></allowwords>
<forbidfloat>1</forbidfloat>
<requirelowestterms>0</requirelowestterms>
<checkanswertype>0</checkanswertype>
<mustverify>0</mustverify>
<options></options>
</input>
<prt>
<name>prt_ans1</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans1</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans1-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The interquartile range is the width of the interval holding the middle 50% of the data, so it measures how spread out the data are.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans1-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The interquartile range is the width of the interval holding the middle 50% of the data, so it measures how spread out the data are.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans2</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans2</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans2-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. &lsquo;Average&rsquo; is a loose word for a typical or central value; the mean and the median are both measures of the centre.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans2-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. &lsquo;Average&rsquo; is a loose word for a typical or central value; the mean and the median are both measures of the centre.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans3</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans3</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans3-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. They are close for symmetric data, but skewness or outliers pull the mean away from the median.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans3-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. They are close for symmetric data, but skewness or outliers pull the mean away from the median.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans4</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans4</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans4-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The range, the maximum minus the minimum, is the simplest measure of variation, although it depends only on the two most extreme values.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans4-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The range, the maximum minus the minimum, is the simplest measure of variation, although it depends only on the two most extreme values.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans5</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans5</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans5-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The standard deviation describes how far, typically, the observations are from the mean.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans5-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The standard deviation describes how far, typically, the observations are from the mean.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans6</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans6</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans6-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The median splits the data into two halves, which is exactly what the second quartile does.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans6-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The median splits the data into two halves, which is exactly what the second quartile does.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans7</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans7</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans7-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The third quartile is found by taking the upper half of the sorted data and finding its median.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans7-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The third quartile is found by taking the upper half of the sorted data and finding its median.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans8</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans8</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans8-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. The IQR ignores the extreme values, so it is the measure of variation to use when the data are skewed or have outliers.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans8-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. The IQR ignores the extreme values, so it is the measure of variation to use when the data are skewed or have outliers.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans9</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans9</sans>
<tans>false</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans9-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; FALSE. The IQR is the <em>third</em> quartile minus the <em>first</em> quartile, \(Q_3 - Q_1\), the range of the middle half of the data.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans9-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is FALSE. The IQR is the <em>third</em> quartile minus the <em>first</em> quartile, \(Q_3 - Q_1\), the range of the middle half of the data.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans10</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans10</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans10-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Three quarters (75%) of the observations lie below the third quartile, which is what the 75th percentile means.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans10-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Three quarters (75%) of the observations lie below the third quartile, which is what the 75th percentile means.</p>]]></text>
</falsefeedback>
</node>
</prt>
<prt>
<name>prt_ans11</name>
<value>1</value>
<autosimplify>1</autosimplify>
<feedbackstyle>1</feedbackstyle>
<node>
<name>0</name>
<answertest>AlgEquiv</answertest>
<sans>ans11</sans>
<tans>true</tans>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
    <truescore>1</truescore>
<truenextnode>-1</truenextnode>
<trueanswernote>prt_ans11-1-T</trueanswernote>
<truefeedback format="html">
<text><![CDATA[<p>Correct &mdash; TRUE. Both are computed directly from the data values (unlike the variance, which is in squared units), so they keep the original units.</p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
    <falsescore>0</falsescore>
<falsenextnode>-1</falsenextnode>
<falseanswernote>prt_ans11-1-F</falseanswernote>
<falsefeedback format="html">
<text><![CDATA[<p>Not quite &mdash; this statement is TRUE. Both are computed directly from the data values (unlike the variance, which is in squared units), so they keep the original units.</p>]]></text>
</falsefeedback>
</node>
</prt>
<qtest>
<testcase>1</testcase>
<description>ans1: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>2</testcase>
<description>ans1: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>false</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans1</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans1-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>3</testcase>
<description>ans2: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>4</testcase>
<description>ans2: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>false</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans2</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans2-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>5</testcase>
<description>ans3: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>6</testcase>
<description>ans3: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>false</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans3</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans3-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>7</testcase>
<description>ans4: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>8</testcase>
<description>ans4: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>false</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans4</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans4-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>9</testcase>
<description>ans5: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>10</testcase>
<description>ans5: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>false</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans5</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans5-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>11</testcase>
<description>ans6: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>12</testcase>
<description>ans6: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>false</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans6</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans6-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>13</testcase>
<description>ans7: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>14</testcase>
<description>ans7: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>false</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans7</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans7-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>15</testcase>
<description>ans8: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>16</testcase>
<description>ans8: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>false</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans8</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans8-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>17</testcase>
<description>ans9: correct (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans9</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans9-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>18</testcase>
<description>ans9: wrong (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>true</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans9</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans9-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>19</testcase>
<description>ans10: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans10</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans10-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>20</testcase>
<description>ans10: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>false</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans10</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans10-1-F</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>21</testcase>
<description>ans11: correct (true)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>true</value>
</testinput>
<expected>
<name>prt_ans11</name>
<expectedscore>1.0000000</expectedscore>
<expectedpenalty>0.0000000</expectedpenalty>
<expectedanswernote>prt_ans11-1-T</expectedanswernote>
</expected>
</qtest>
<qtest>
<testcase>22</testcase>
<description>ans11: wrong (false)</description>
<testinput>
<name>ans1</name>
<value>true</value>
</testinput>
<testinput>
<name>ans2</name>
<value>true</value>
</testinput>
<testinput>
<name>ans3</name>
<value>true</value>
</testinput>
<testinput>
<name>ans4</name>
<value>true</value>
</testinput>
<testinput>
<name>ans5</name>
<value>true</value>
</testinput>
<testinput>
<name>ans6</name>
<value>true</value>
</testinput>
<testinput>
<name>ans7</name>
<value>true</value>
</testinput>
<testinput>
<name>ans8</name>
<value>true</value>
</testinput>
<testinput>
<name>ans9</name>
<value>false</value>
</testinput>
<testinput>
<name>ans10</name>
<value>true</value>
</testinput>
<testinput>
<name>ans11</name>
<value>false</value>
</testinput>
<expected>
<name>prt_ans11</name>
<expectedscore>0.0000000</expectedscore>
<expectedpenalty>0.1000000</expectedpenalty>
<expectedanswernote>prt_ans11-1-F</expectedanswernote>
</expected>
</qtest>
</question>
</quiz>